# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

```python
#import necessary packages 
import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd
```

```python
#find the database 
host = 'slucor2022-instance-1.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2022'
password = 'SLUcor2022'
database = 'hds5210'
# database = ''
```

```python
conn = create_engine('mysql+pymysql://' + 
                     user + ':' + 
                     password + '@' + 
                     host + '/' + 
                     database, echo=False)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
#name my table in the database according to username 
import getpass
myname = getpass.getuser().split('-')[1]
myname
MYTABLE=myname + '_data'
```

```python
#do this so code can be re-run 
with conn.connect() as c:
    c.execute('DROP TABLE IF EXISTS ' + MYTABLE)
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'lillyeversman')
assert(conn.connect())
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

https://opendata.arcgis.com/api/v3/datasets/ed0bc689bde246b18835c7529ba4efb4_0/downloads/data?format=csv&spatialRefId=4326

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
#bring in the data from the link and reset index so columns are as we need them to be 
data = pd.read_csv("EHR_eligable.csv")
data.reset_index()
```

Output:
```
index           X          Y  OBJECTID  \
0        0 -124.142009  40.783559         1   
1        1 -122.086674  37.632915         2   
2        2 -122.295861  38.325402         3   
3        3 -120.766467  38.350920         4   
4        4 -122.435820  37.769049         5   
..     ...         ...        ...       ...   
318    318 -122.173214  37.436328       319   
319    319 -118.186632  33.808293       320   
320    320 -121.455450  38.553012       321   
321    321 -118.286401  34.067178       322   
322    322 -117.152024  32.798322       323   

                                         Provider_Name         NPI    CCN  \
0             ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1609858950  50006   
1                             HAYWARD SISTERS HOSPITAL  1942298153  50002   
2             ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1235218785  50009   
3                              SUTTER VALLEY HOSPITALS  1447494323  50014   
4                                 Sutter Bay Hospitals 
```

```python
data1 = data.reset_index()
```

```python
#read in using sql query show tables function 
data1.to_sql(MYTABLE, conn, index=False) 
data2 = pd.read_sql_query("SHOW TABLES", conn)
data2
```

Output:
```
Tables_in_hds5210
0          asahu123_data
1                   data
2          ecdmdphd_data
3              ed_visits
4        havenaspen_data
5         havenaspen_ipv
6      hospital_services
7              hospitals
8     lillyeversman_data
9          paulboal_data
10     pbstudent_HospPay
11   pbstudent_PersonPay
12     pbstudent_ZipCode
13        pbstudent_data
14            population
15           rasalt_data
16       rasalt_data_tmp
17      rasalt_data_tmp1
18      rasalt_data_tmp2
19           shiv80_data
20       svelicheti_data
21   system_affiliations
22  total_per_population
23         vikas060_data
24          vikas060data
25  vineethmunigati_data
```

```python
#create tables so you can see tables utilized in assignemnt 
tables = {}

for n,t in pd.read_sql_query("SHOW TABLES", conn).iterrows():
    name = str(t['Tables_in_hds5210'])
    print(name)
    tables[name] = pd.read_sql("SELECT * FROM "+name,conn)
```

Output:
```
asahu123_data
data
ecdmdphd_data
ed_visits
havenaspen_data
havenaspen_ipv
hospital_services
hospitals
lillyeversman_data
paulboal_data
pbstudent_HospPay
pbstudent_PersonPay
pbstudent_ZipCode
pbstudent_data
population
rasalt_data
rasalt_data_tmp
rasalt_data_tmp1
rasalt_data_tmp2
shiv80_data
svelicheti_data
system_affiliations
total_per_population
vikas060_data
vikas060data
vineethmunigati_data
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
lil = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
lil.head()
```

Output:
```
index           X          Y  OBJECTID  \
0      0 -124.142009  40.783559         1   
1      1 -122.086674  37.632915         2   
2      2 -122.295861  38.325402         3   
3      3 -120.766467  38.350920         4   
4      4 -122.435820  37.769049         5   

                              Provider_Name         NPI    CCN  \
0  ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1609858950  50006   
1                  HAYWARD SISTERS HOSPITAL  1942298153  50002   
2  ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1235218785  50009   
3                   SUTTER VALLEY HOSPITALS  1447494323  50014   
4                      Sutter Bay Hospitals  1659439834  50008   

  Business_Street_Address  Business_City Business_County  ...  \
0         2700 Dolbeer St         Eureka        Humboldt  ...   
1      27200 Calaroga Ave        Hayward         Alameda  ...   
2         1000 Trancas St           Napa            Napa  ...   
3        200 Mission Blvd        Jackson          Amador  ...   
4       CA
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (323,22))
assert(list(dxyz.columns) == ['index', 'X', 'Y', 'OBJECTID', 'Provider_Name', 'NPI', 'CCN',
       'Business_Street_Address', 'Business_City', 'Business_County',
       'Business_ZIP_Code', 'Business_State_Territory', 'Payment_Year_Number',
       'Program_Type', 'Medicaid_EP_Hospital_Type', 'total_payments',
       'Last_Payment_Criteria', 'Recent_Disbursement_Amount', 'Latitude',
       'Longitude', 'Last_Program_Year', 'Last_Payment_Year'])
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Business_ZIP_Code` field and `population`'s `Zip` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

```python
#make tables into objects so we can merge with pd 
population = tables['population']
```

```python
mytable = tables['lillyeversman_data']
```

```python
#merge tables on zip code fields and sort in order of population so we can specify which zip code has the highest population
answer1 = pd.merge(mytable, population, left_on = 'Business_ZIP_Code', right_on = 'Zip', how = 'inner')
answer2= answer1.sort_values(by ='Population' , ascending=False)
answer2.head()
```

Output:
```
index_x           X          Y  OBJECTID  \
280      275 -118.099150  33.912613       276   
103       74 -117.432081  34.071840        75   
46        21 -118.445983  34.324488        22   
77        48 -117.439101  33.926868        49   
234      217 -117.685655  34.026272       218   

                         Provider_Name         NPI    CCN  \
280       CPH HOSPITAL MANAGEMENT, LLC  1194016923  50771   
103        KAISER FOUNDATION HOSPITALS  1356496772  50140   
46               COUNTY OF LOS ANGELES  1134143415  50040   
77   DOCTORS HOSPITAL OF RIVERSIDE LLC  1124596382  50102   
234      VERITAS HEALTH SERVICES, INC.  1962407460  50586   

    Business_Street_Address Business_City Business_County  ...  \
280     13100 Studebaker Rd       Norwalk     Los Angeles  ...   
103      9961 Sierra Avenue       Fontana  San Bernardino  ...   
46   14445 Olive View Drive        Sylmar     Los Angeles  ...   
77          3865 Jackson St     Riverside       Riverside  ...   
234         5
```

```python
mytable.head()
```

Output:
```
index           X          Y  OBJECTID  \
0      0 -124.142009  40.783559         1   
1      1 -122.086674  37.632915         2   
2      2 -122.295861  38.325402         3   
3      3 -120.766467  38.350920         4   
4      4 -122.435820  37.769049         5   

                              Provider_Name         NPI    CCN  \
0  ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1609858950  50006   
1                  HAYWARD SISTERS HOSPITAL  1942298153  50002   
2  ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1235218785  50009   
3                   SUTTER VALLEY HOSPITALS  1447494323  50014   
4                      Sutter Bay Hospitals  1659439834  50008   

  Business_Street_Address  Business_City Business_County  ...  \
0         2700 Dolbeer St         Eureka        Humboldt  ...   
1      27200 Calaroga Ave        Hayward         Alameda  ...   
2         1000 Trancas St           Napa            Napa  ...   
3        200 Mission Blvd        Jackson          Amador  ...   
4       CA
```

```python
#select the row with the top value in my sorted list - row 280 and print the value of NPI 
answer3 = answer2['NPI'][280].tolist()
answer = [answer3]
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == [1194016923])
```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

## First using SQL

Please see below sql area for pandas solution - much better!!!

```python
#read in the data and see total payments 
pd.read_sql_query("SELECT NPI, Medicaid_EP_Hospital_Type, Business_ZIP_Code,total_payments, SUM(total_payments) AS total FROM lillyeversman_data GROUP BY Business_ZIP_Code, Medicaid_EP_Hospital_Type ", conn)
```

Output:
```
NPI Medicaid_EP_Hospital_Type  Business_ZIP_Code  total_payments  \
0    1639166200      Acute Care Hospitals              90004      1093021.73   
1    1881656460      Acute Care Hospitals              90012      3614739.93   
2    1114081056      Acute Care Hospitals              90015     10045484.04   
3    1508859323      Acute Care Hospitals              90017      3899845.21   
4    1508908591      Children's Hospitals              90020       427640.22   
..          ...                       ...                ...             ...   
262  1801973052      Acute Care Hospitals              96104       184615.39   
263  1326044140      Acute Care Hospitals              96122       531907.08   
264  1538197231      Acute Care Hospitals              96130      1371962.88   
265  1003867565      Acute Care Hospitals              96150       786769.13   
266  1538265780      Acute Care Hospitals              96161       621315.39   

           total  
0     1093021.73  
1     3614739
```

```python
sql = pd.read_sql_query("SELECT NPI, Medicaid_EP_Hospital_Type, Business_ZIP_Code,total_payments, SUM(total_payments) AS total FROM lillyeversman_data GROUP BY Business_ZIP_Code ORDER BY total DESC", conn)
```

```python
#merge the table with popualtion table to calculate the per capital totals
avgcapital = pd.read_sql_query("""
  SELECT NPI, Medicaid_EP_Hospital_Type, Business_ZIP_Code,total_payments, population, SUM(total_payments) AS total
  FROM 
    lillyeversman_data h INNER JOIN
    population v ON h.Business_ZIP_Code = v.Zip
  GROUP BY Business_ZIP_Code ORDER BY total DESC
""", conn)
```

```python
avgcapital['Avg_Pay_per_Capital'] = avgcapital['total']/avgcapital['population']
avgcapital
```

Output:
```
NPI Medicaid_EP_Hospital_Type  Business_ZIP_Code  total_payments  \
0    1770659336      Acute Care Hospitals              94107       309782.57   
1    1861407637      Acute Care Hospitals              90806      1110809.22   
2    1215927470      Acute Care Hospitals              90033      6017613.32   
3    1194731935      Acute Care Hospitals              92354      7465049.37   
4    1912982216      Acute Care Hospitals              92868      2221862.07   
..          ...                       ...                ...             ...   
251  1801973052      Acute Care Hospitals              96104       184615.39   
252  1710116116      Acute Care Hospitals              95688       179195.17   
253  1194895227      Acute Care Hospitals              94903       175861.23   
254  1386714814      Acute Care Hospitals              94063       126249.13   
255  1740354851      Acute Care Hospitals              95337        68865.12   

     population        total  Avg_Pay_per_Capital  
```

```python
#find averages
import numpy as np
avgcapital2 = pd.pivot_table(data = avgcapital, index = ['Medicaid_EP_Hospital_Type'], values = 'Avg_Pay_per_Capital', aggfunc = np.mean)
avgcapital2
```

Output:
```
Avg_Pay_per_Capital
Medicaid_EP_Hospital_Type                     
Acute Care Hospitals                110.378444
Children's Hospitals                468.554016
```

## Now using Pandas instead of SQL


I MOVED ASSERTIONS BELOW BECAUSE PANDAS WAS GETTING THE CORRECT VALUES TO TEST WITH ASSERTIONS. NOT SURE WHY SQL IS AVERAGING DIFFERENTLY! 

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
#install packages and import data 
import numpy as np
test = pd.pivot_table(data = mytable, index = ['Business_ZIP_Code'], values = 'total_payments', columns = ['Medicaid_EP_Hospital_Type'], aggfunc = np.sum,   margins = True, margins_name='Total')
```

```python
#sort by totals to see which zip code has highest values
test.sort_values(by = 'Total', ascending=False)
```

Output:
```
Medicaid_EP_Hospital_Type  Acute Care Hospitals  Children's Hospitals  \
Business_ZIP_Code                                                       
Total                              7.399908e+08           52879206.24   
94107                              4.677515e+07                   NaN   
90806                              1.502282e+07            7951873.30   
90033                              1.980909e+07                   NaN   
92354                              1.692635e+07                   NaN   
...                                         ...                   ...   
96104                              1.846154e+05                   NaN   
95688                              1.791952e+05                   NaN   
94903                              1.758612e+05                   NaN   
94063                              1.262491e+05                   NaN   
95337                              6.886512e+04                   NaN   

Medicaid_EP_Hospital_Type         Total  
Business
```

```python
#groupresults by zip code and hospital type and use sum function
myresults = mytable.groupby(['Business_ZIP_Code', 'Medicaid_EP_Hospital_Type']).sum()
```

```python
#group results and sort them, reset index to get the index as column values so we can join below
myresults.head()
grouped = myresults.sort_values(by ='total_payments' , ascending=False).reset_index()
grouped
```

Output:
```
Business_ZIP_Code Medicaid_EP_Hospital_Type  index            X  \
0                94107      Acute Care Hospitals   2333 -2570.234116   
1                90033      Acute Care Hospitals    461  -354.628262   
2                92354      Acute Care Hospitals    417  -234.522344   
3                90806      Acute Care Hospitals    836  -472.759362   
4                95128      Acute Care Hospitals     97  -243.869114   
..                 ...                       ...    ...          ...   
262              96104      Acute Care Hospitals    287  -120.171848   
263              95688      Acute Care Hospitals    271  -121.938457   
264              94903      Acute Care Hospitals    197  -122.554103   
265              94063      Acute Care Hospitals    203  -122.222744   
266              95337      Acute Care Hospitals    262  -121.246101   

              Y  OBJECTID          NPI      CCN  Payment_Year_Number  \
0    793.311023      2354  31304774953  1056183                   76
```

```python
#use merge to join the zip code fields
together = pd.merge(grouped,population, left_on = 'Business_ZIP_Code', right_on = 'Zip',how = 'inner')
together
```

Output:
```
Business_ZIP_Code Medicaid_EP_Hospital_Type  index_x            X  \
0                94107      Acute Care Hospitals     2333 -2570.234116   
1                90033      Acute Care Hospitals      461  -354.628262   
2                92354      Acute Care Hospitals      417  -234.522344   
3                90806      Acute Care Hospitals      836  -472.759362   
4                90806      Children's Hospitals      319  -118.186632   
..                 ...                       ...      ...          ...   
256              96104      Acute Care Hospitals      287  -120.171848   
257              95688      Acute Care Hospitals      271  -121.938457   
258              94903      Acute Care Hospitals      197  -122.554103   
259              94063      Acute Care Hospitals      203  -122.222744   
260              95337      Acute Care Hospitals      262  -121.246101   

              Y  OBJECTID          NPI      CCN  Payment_Year_Number  \
0    793.311023      2354  31304774953  1056
```

```python
#create new calculated column for per person charges 
together['Avg_Pay_per_Capita'] = together.total_payments/ together.Population
together
```

Output:
```
Business_ZIP_Code Medicaid_EP_Hospital_Type  index_x            X  \
0                94107      Acute Care Hospitals     2333 -2570.234116   
1                90033      Acute Care Hospitals      461  -354.628262   
2                92354      Acute Care Hospitals      417  -234.522344   
3                90806      Acute Care Hospitals      836  -472.759362   
4                90806      Children's Hospitals      319  -118.186632   
..                 ...                       ...      ...          ...   
256              96104      Acute Care Hospitals      287  -120.171848   
257              95688      Acute Care Hospitals      271  -121.938457   
258              94903      Acute Care Hospitals      197  -122.554103   
259              94063      Acute Care Hospitals      203  -122.222744   
260              95337      Acute Care Hospitals      262  -121.246101   

              Y  OBJECTID          NPI      CCN  Payment_Year_Number  \
0    793.311023      2354  31304774953  1056
```

```python
#make a new df with just specified columns at hand then sort in descending order
newdf = together[['Medicaid_EP_Hospital_Type', 'Avg_Pay_per_Capita', 'Zip']].copy()
newdf.sort_values(by ='Avg_Pay_per_Capita', ascending=False)
```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita    Zip
0        Acute Care Hospitals         1574.496706  94107
41       Children's Hospitals         1282.361565  94304
170      Acute Care Hospitals         1065.544770  96028
15       Acute Care Hospitals         1043.182013  93721
51       Acute Care Hospitals          851.476203  95231
..                        ...                 ...    ...
257      Acute Care Hospitals            4.701558  95688
255      Acute Care Hospitals            4.225844  91367
249      Acute Care Hospitals            4.132765  91706
259      Acute Care Hospitals            3.496237  94063
260      Acute Care Hospitals            1.693807  95337

[261 rows x 3 columns]
```

```python
#make pivot table to show the aggregate average function by hopsital type 
newdf2 = pd.pivot_table(data = newdf, index = ['Medicaid_EP_Hospital_Type'], values = 'Avg_Pay_per_Capita', aggfunc = np.mean) 
newdf2
```

Output:
```
Avg_Pay_per_Capita
Medicaid_EP_Hospital_Type                    
Acute Care Hospitals               106.582669
Children's Hospitals               329.550946
```

```python
#rename as answer so assertions can be tested 
answer = newdf2.reset_index()
answer
```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita
0      Acute Care Hospitals          106.582669
1      Children's Hospitals          329.550946
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 106.583)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 329.551)

```

```python
round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3)
```

Output:
```
106.583
```



---



## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git pull
    !git add week15_assignment.ipynb
    !git commit -a -m "Submitting the week 15 programming assignment"
    !git push
else:
    print('''
    
OK. We can wait.
''')
    
```

```python

```